---
id: 66b205e6eacba4c4e54ea434
title: Build an Email Masker
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

In this lab, you will mask the username part of an email address with asterisks. Masking is a term used to hide or replace sensitive information with asterisks or other characters.

For example, if the email address was `myEmail@email.com`, then the masked email address will be `m*****l@email.com`.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. Create a function named `maskEmail` that takes `email` as an argument.
2. Inside the function, you should mask the `email` and append the domain name to it. Remember that you can use methods like `slice`, `repeat`, `indexOf` or even `replace` to help you.
3. Outside the function, declare a variable named `email` to store the email address you want to mask.
4. Call the `maskEmail` function with the `email` variable and output the result to the console. 
5. `maskEmail("apple.pie@example.com")` should return `"a*******e@example.com"`.
6. `maskEmail("freecodecamp@example.com")` should return `"f**********p@example.com"`.
7. `maskEmail("info@test.dev")` should return `"i**o@test.dev"`.
8. `maskEmail("user@domain.org")` should return `"u**r@domain.org"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

You should define a function named `maskEmail`.

```js
assert.isFunction(maskEmail);
```

Your `maskEmail` function should take one argument.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.functions.maskEmail.parameters, 1);
```

You should have an `email` variable outside the function.

```js
assert.isDefined(email);
```

You should assign a valid email address to your `email` variable.

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

Your `maskEmail` function should return `"a*******e@example.com"` when called with `"apple.pie@example.com"`.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

Your `maskEmail` function should return `"f**********p@example.com"` when called with `"freecodecamp@example.com"`.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

Your `maskEmail` function should return `"i**o@test.dev"` when called with `"info@test.dev"`.

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

Your `maskEmail` function should return `"u**r@domain.org"` when called with `"user@domain.org"`.

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

Your `maskEmail` function should produce the correct result for other emails, such as `"johnappleseed@email.com"`.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

You should log the output of calling `maskEmail` with `email` as argument.

```js
assert.deepInclude(spy.calls, [maskEmail(email)]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
